

# Homework 1: Basic Python, Arrays, and DataFrames

## Due Monday, April 13th at 11:59PM

Welcome to Homework 1! This week's homework will cover basic Python, arrays, and DataFrames. You can find additional help on these topics in [Chapter 1](https://www.inferentialthinking.com/chapters/01/what-is-data-science.html) of Computational and Inferential Thinking and [BPD 1-11](https://notes.dsc10.com/01-getting_started/tools.html) in the `babypandas` notes.


### Instructions

Remember to start early and submit often. You are given six slip days throughout the quarter to extend deadlines. See the syllabus for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

**Important**: For homeworks, the `otter` tests don't usually tell you that your answer is correct. More often, they help catch careless mistakes. It's up to you to ensure that your answer is correct. If you're not sure, ask someone (not for the answer, but for some guidance about your approach). These are great questions for office hours (the schedule can be found [here](https://dsc10.com/calendar)) or Campuswire. Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

```python
# Please don't change this cell, but do make sure to run it.
import babypandas as bpd
import numpy as np
np.set_printoptions(legacy='1.21')

import matplotlib.pyplot as plt
plt.style.use('seaborn-v0_8-colorblind')
plt.rcParams['figure.figsize'] = (10, 5)

import otter
grader = otter.Notebook("hw1.ipynb")
```

## 1. Characters in *The Wonderful Wizard of Oz* 👠🧠❤️🦁

In Lecture 1, we counted the number of times that the characters Amy, Beth, Jo, Meg, and Laurie were named in each chapter of the classic book _Little Women_. In this question, we'll look at another classic book – _The Wonderful Wizard of Oz (1900)_ by L. Frank Baum and illustrated by W.W. Denslow. In 1956, the copyright protections of the original book (but not the 1939 film produced by Metro-Goldwyn-Mayer!) expired, and so sites like Project Gutenberg are now able to post copies of [the book](https://www.gutenberg.org/files/55/55-h/55-h.htm) without violating any copyright laws. 

<center><img src=images/wizard-of-oz-illustration.png width=350><br>(<a href="https://www.nypl.org/events/exhibitions/galleries/childhood/item/5458">source</a>)</center>

Four of the main characters in _The Wonderful Wizard of Oz_ are Dorothy Gale (👠), Scarecrow (🧠), The Tin Man (❤️), and The Cowardly Lion (🦁).

Below, we've written code that shows the number of mentions of each of these four characters in the first 12 chapters of the book. However, instead of displaying this information in a line chart, as was done in Lecture 1, we will use a bar chart.

Run the cell below to generate this bar chart of the first 12 chapters. This cell contains code that hasn't yet been covered in the course. It isn't expected that you'll understand the code, but you should be able to interpret the bar chart it generates. All questions can be answered by looking at the bar chart alone!

```python
# Open the book and split it into chapters.
book_file = 'data/the-wonderful-wizard-of-oz.txt'
raw_book = open(book_file, encoding="utf-8").read()
end_pos = raw_book.index('Chapter XIII')
chapters = raw_book[:end_pos].split('Chapter ')[1:]

# Find the number of words in each chapter.
chapter_lengths = (np.array([len(c.split(' ')) for c in chapters]) / 100)

# Find the number of mentions per 100 words for each character and chapter.
characters = bpd.DataFrame().assign(
    Chapter=np.arange(1, 13),
    Dorothy=np.char.count(chapters, 'Dorothy') / chapter_lengths,
    Scarecrow=np.char.count(chapters, 'Scarecrow') / chapter_lengths,
    Lion=np.char.count(chapters, 'Lion') / chapter_lengths,
    Tin_Man=np.char.count(chapters, 'Tin Woodman') / chapter_lengths
)

characters.plot(kind='bar', x='Chapter', figsize=(14, 8))
plt.ylabel('Mentions per 100 words in chapter');
```

The height of the bar for Dorothy in Chapter 1 is 1.62 and the y-axis of this graph is “Mentions per 100 words in chapter”. This means that 1.62 of every 100 words in Chapter 1 are `"Dorothy"`. In other words, $1.62\%$ of the words in Chapter 1 are `"Dorothy"`.

**Question 1.1.**  The following sentences are the very first time Scarecrow, Tin Man (real name Tin Woodman), and the Lion are mentioned in the book, respectively:

> There was a great cornfield beyond the fence, and not far away she saw a Scarecrow, placed high on a pole to keep the birds from the ripe corn.

> "Oil my neck, first," replied the Tin Woodman.

> Just as he spoke there came from the forest a terrible roar, and the next moment a great Lion bounded into the road.

Based on the above bar chart, which chapters are each of these sentences from? Assign the variables `scarecrow_debut`, `tin_man_debut`, `lion_debut` to an integer between 1 and 12.

```python
scarecrow_debut = ...
tin_man_debut = ...
lion_debut = ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** Dorothy is mentioned 38 times in Chapter 12. How many times is Lion mentioned in Chapter 12? Assign the variable `lion_mentioned` to 1, 2, 3, 4, or 5 corresponding to the answer choices below. Remember to input your answer as a **number from 1-5 corresponding to the option numbers below, and not to the number of times Lion is mentioned.** For example, if you think Lion is mentioned 13 times, you should set `lion_mentioned` to 1, not 13.

1. 13 times
1. 15 times
1. 17 times
1. 19 times
1. 20 times

```python
lion_mentioned = ...
```

```python
grader.check("q1_2")
```

**Question 1.3.** Chapter 11 is quite long, as it is the chapter when the characters reach the Emerald City. For reference, Chapter 8 is 1764 words long while Chapter 11 is 3294 words long. Based on this information and the plot, which of the following is correct? Assign the variable `oz_part3` to 1, 2, 3, or 4.

1. Scarecrow is mentioned more times in Chapter 8 than in Chapter 11. Tin Man is mentioned more times in Chapter 8 than in Chapter 11.
1. Scarecrow is mentioned more times in Chapter 8 than in Chapter 11. Tin Man is mentioned fewer times in Chapter 8 than in Chapter 11.
1. Scarecrow is mentioned fewer times in Chapter 8 than in Chapter 11. Tin Man is mentioned more times in Chapter 8 than in Chapter 11.
1. Scarecrow is mentioned fewer times in Chapter 8 than in Chapter 11. Tin Man is mentioned fewer times in Chapter 8 than in Chapter 11.

```python
oz_part3 = ...
```

```python
grader.check("q1_3")
```

**Question 1.4.** Which of the following is a valid conclusion we can make based off of the above plot alone? Assign `oz_part4` to 1, 2, 3, 4, or 5. There is only one correct answer.

1. Lion is mentioned more times in Chapter 10 than he is in Chapter 6.

1. Chapter 12 is the chapter in which Tin Man is mentioned the most.

1. Lion is mentioned roughly the same number of times in each of Chapters 6, 8, and 10.

1. Chapter 8 is the chapter in which Dorothy is mentioned the least.

1. Scarecrow and Tin Man are mentioned roughly the same number of times in Chapter 6.

```python
oz_part4 = ...
```

```python
grader.check("q1_4")
```

**_Note:_**  The tests in this section only check that you set each variable to a number in the correct range. Unlike in labs, tests in homeworks **do not** check that you answered correctly; they only check that your answer is *reasonable*, or in the correct format. To put it another way: all of your tests might pass, but that doesn't mean you'll get full credit – some of your answers may still be wrong. It's up to you to make sure that they're right!

## 2. Python Basics 🐍

**Question 2.1.** When you run the following cell, Python produces a cryptic error message.

```python
2024 = 2025 - 1.0
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `basics_1` below to indicate your answer.

1. This is trying to create a variable named `2024`, which doesn't make sense because `2024` is a number.

1. The left hand side is an `int`, while the right hand side is a `float`. The left side should be `2024.0` instead.

1. The result should be written after the calculation. It should be `2025 - 1.0 = 2024`.

1. Python is not able to subtract a `float` from an `int` because they are of different data types.


<font color=red>**🚨 Important: Once you have finished this question, "comment"  out the above code cell out by replacing it with `# 2024= 2025 - 1.0`. This will prevent the error message from appearing when your notebook is graded.**</font>

***Note:*** A shortcut for "commenting" out code is to highlight the code and press `command` or `control` and `/`.

```python
basics_1 = ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** Consider the following poorly-written code.

```python
three = 3
four = 4
three = three + four
three = three * four
three = three * three
three = -three
```

As this code executes, what values does the variable `three` take on? Assign 1, 2, 3, or 4 to `basics_2` to indicate your answer.

1. The variable `three` takes on the values 3, 7, 12, 144, -144.

1. The variable `three` takes on the values 3, 7, 28, 784, -784.

1. The variable `three` takes on the values 3, 7, 28, 84, -84.

1. The variable `three` takes on the values 3, 7, 12, 36, -36.

```python
basics_2 = ...
```

```python
grader.check("q2_2")
```

## 3. Road Trip   🚘 

Over spring break, you plan on going on a road trip to Palm Springs with your friends Abel and Oren. Before you go, you want to plan whose car (Abel's or Oren's) and which route (Route A or Route B) to take. Answer the questions below using Python to perform all the intermediate calculations such as adding, squaring, and dividing.

<font color=red>**🚨 Important: The `math` package has not been imported. You don't need it for this question, and you should not import it, otherwise the Gradescope autograder may error.**</font>

**Question 3.1.** First, you need to decide on the route. You plan to take the route for which the average speed is faster. 

For Route A, you will have to take 3 freeways, each with a different speed limit. Below is a table showing the speed limits (which we will assume is the speed you will travel at) and the time you will have to spend on each freeway on Route A.

| Freeway | Speed Limit (miles per hour) | Time (hours)|
| --- | --- | --- |
| I-5 | 70 | 3 |
| CA-73 | 55 | 2 |
| CA-133 | 45 | 1 |

Using this information, calculate the average speed, in miles per hour, if you take Route A, and assign your answer to the variable `route_A`. Recall from math and physics that the average speed is the total distance driven divided by the total time taken.

```python
# Feel free to define intermediate variables to use in your solution.
total_distance = ...
total_time = ...
route_A = ...
route_A
```

```python
grader.check("q3_1")
```

**Question 3.2.** Next, let's consider the other route you might take. For Route B, you will have to take 3 freeways, each with a different speed limit. Below is a table showing the speed limits (which we will assume is the speed you will travel at) and the *distance* you will need to travel on each freeway on Route B.

| Freeway | Speed Limit (miles per hour) | Distance (miles)|
| --- | --- | --- |
| I-5 | 70 | 150 |
| CA-57 | 65 | 110 |
| CA-109 | 55 | 90 |

Using this information, calculate the average speed, in miles per hour, if you take Route B, and assign your answer to the variable `route_B`.

Note that the third column is **`'Distance (miles)'`, not `'Time (hours)'`**. Unlike in Question 3.1, you aren't given the amount of time you'll spend on each freeway; you need to compute these times yourself. To calculate the time it will take on each freeway, divide the distance for that freeway by the speed for that freeway. Finally, add up the times for each freeway to find the total time.

```python
# Feel free to define intermediate variables to use in your solution.
total_distance2 = ...
total_time2 = ...

route_B = ...
route_B
```

```python
grader.check("q3_2")
```

**Question 3.3.** Now it's time to decide whose car to take (Abel's or Oren's). You decide to take the car with the biggest fuel tank. 

Unfortunately, Abel doesn't know the exact volume of his car's fuel tank, only that the tank has a height of 16 inches, a width of 12 inches, and a length of 32 inches. Oren doesn't know the exact volume of his car's fuel tank either, but knows that his tank is in the shape of a cube. What is the length of one of the sides of Oren's fuel tank (in inches) so that it stores the same amount of fuel as Abel's? Save your answer in the variable `side_length`.

```python
# Feel free to define intermediate variables to use in your solution.
side_length = ...
side_length
```

```python
grader.check("q3_3")
```

In this problem, though you calculated three different quantities in three different ways, all of your results are actually considered **means**, of various kinds!

In Question 3.1, given $n$ values $x_1, x_2, ..., x_n$, you found an *arithmetic mean*, using the formula

$${x_1+x_2+...+x_n \over n},$$

where the numerator represented total distance and the denominator represented total time. An arithmetic mean is the usual type of mean or average you're used to seeing. It turns out that you actually computed a more sophisticated arithmetic mean, known as a _weighted arithmetic mean_, 

$$\frac{w_1 x_1 + w_2 x_2 + ... + w_n x_n}{w_1 + w_2 + ... + w_n},$$

where the weights $w_1, w_2, w_3$ were the times on each freeway.

In Question 3.2, given  $n$ values $x_1, x_2, ..., x_n$, you found a *harmonic mean*, using the formula

$${n \over {{1 \over x_1}+{1 \over x_2}+ ... + {1 \over x_n}}},$$ 


where the numerator represented total distance and the denominator represented total time. To calculate the total time, you needed to sum the time taken for each part of the trip, calculated using the fact that time is distance over speed. Again, it turns out that you actually computed the _weighted harmonic mean_, but this time the weights were the distances on each of the freeways. If you're curious, see the formula [here](https://en.wikipedia.org/wiki/Harmonic_mean#Weighted_harmonic_mean).

Finally in Question 3.3, given $n$ values $x_1, x_2, ..., x_n$, you found a *geometric mean*, using the formula 

$${\sqrt[n]{x_1 \cdot x_2 \cdot ... \cdot x_n}},$$ 

where each value represented a dimension of the fuel tank. 

As you can see, there are many different notions of the mean. You'll learn about some of them if you take DSC 40A!

## 4. Social Media Usage Among Communities 📱

In this problem, we want to compare and contrast social media usage among Americans who live in different types of communities (urban, suburban, and rural). The data below comes from [Pew Research Center's 2023 Social Media Fact Sheet](https://www.pewresearch.org/internet/fact-sheet/teens-and-social-media-fact-sheet/?tabId=tab-26288fdf-8b5e-4076-97e4-84f6ce9ecbf7). It shows what percent of residents of each type of community use TikTok, Snapchat, and Instagram.

| Percent of community using platform       | Urban  | Suburban  | Rural |
|-------------------------------------------|-------------|-----------|-----------|
| TikTok        | 67 | 62 | 62 | 
| Snapchat       | 59 | 58 | 64 | 
| Instagram    | 65 | 58  | 55 | 

We define the **dissimilarity** between two communities as the largest absolute difference between their three respective usage percentages.

To better understand dissimilarity, consider the following hypothetical situation, where we compare the social media usage in community A and community B. Suppose:
* Community A's *usage of Facebook* is **10 percent more** than community B's.
* Community A's *usage of Instagram* is **4 percent less** than community B's.
* Community A's *usage of LinkedIn* is **7 percent less** than community B's.

Then, we would say the dissimilarity between community A and community B is 10, since 10 is larger than both 4 and 7.

**Question 4.1.** 
Using this method, compute the dissimilarity between urban and suburban communities.  Assign the result to the variable `dissimilarity`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting) rather than calculating the expression yourself. 

**_Hint:_**  The built-in `abs` function computes absolute values. 

```python
dissimilarity = ...
dissimilarity
```

```python
grader.check("q4_1")
```

**Question 4.2.** Which pair of communities is **most** dissimilar, according to this measurement? Assign either 1, 2, or 3 to the variable `most_dissimilar` below. Check whether your answer matches up with your intuition.

1. urban and suburban communities
1. suburban and rural communities
1. urban and rural communities

```python
most_dissimilar = ...
```

```python
grader.check("q4_2")
```

**Question 4.3.** It turns out that if we were to eliminate a certain one of the three social media percentages in the table (getting rid of one of the rows) and recalculate dissimilarities based on the remaining two percentages only, we would find the dissimilarity between each pair of communities to be the same as if we had used all three percentages. In other words, one of the three rows of the table ends up not factoring into the calculation for dissimilarity, for all three pairs of communities.

Which percentage can be eliminated without changing the dissimilarity of any pair of communities in the table? Assign either 1, 2, or 3 to the variable `disposable` below.

1. TikTok
1. Snapchat
1. Instagram

```python
disposable = ...
```

```python
grader.check("q4_3")
```

## 5. Arrays 🗃️

**Question 5.1.** Make an array called `quirky_numbers` containing the following numbers (in the given order):

1. The square root of 67
2. 33 degrees, in radians
3. $\frac{3^4 + 9^4}{1000}$
4. The mathematical constant $e$ divided by 7 (that is, $\frac{e}{7}$)
5. The base 10 logarithm of 88

<font color=red>**🚨 Important: The `math` package has not been imported. You don't need it for this question, and you should not import it, otherwise the Gradescope autograder may error.**</font>

**_Hint:_**  Check out the functions and constants available in the `numpy` module, which has been imported as `np`. If you're unsure of what function to use, a quick Google search should do the trick. Do **not** round any of the values in the array. 

**_Note:_**  In this problem, as with all others, we'll only check that your answer is correct. There may be several valid ways to produce the correct answer.

```python
quirky_numbers = ...
quirky_numbers
```

```python
grader.check("q5_1")
```

**Question 5.2.** Make an array called `likes` containing the following three strings, in the specified order:
- `'I like surfing'`
- `'otters'`
- `'and popcorn.'`



```python
likes = ...
likes
```

```python
grader.check("q5_2")
```

<center><img src=images/otter.webp width=300>
<br>
<a href="https://www.npr.org/2023/07/13/1187295769/otter-santa-cruz-surfboard-surfers-california">An otter in Santa Cruz is hassling surfers — and stealing their boards</a></center>


In [Lecture 3](https://dsc10.com/resources/lectures/lec03/lec03.html#String-methods), we looked at several string methods, like `upper` and `replace`. Strings have another method that we haven't seen yet, called `join`. `join` takes one argument, an array of strings, and it returns a single string. Specifically, `some_string.join(some_array)` evaluates to a new string consisting of all of elements in `some_array`, with `some_string` inserted **in between each element**.

For example, `'-'.join(np.array(['call', '858', '534', '2230']))` evaluates to `'call-858-534-2230'`.

**Question 5.3.** Use the array `likes` and the method `join` to make two strings:

1. `'I like surfing, otters, and popcorn.'` (call this one `by_comma`)
1. `'I like surfing otters and popcorn.'` (call this one `by_space`)

```python
by_comma = ...
by_space = ...

# Don't change the lines below.
print(by_comma)
print(by_space)
```

```python
grader.check("q5_3")
```

Now let's get some practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *integer position*, with the position of the first element being zero. That's probably not the way you learned to count, so it's easy to get mixed up here. Be careful!

**Question 5.4.** The cell below creates an array of strings.

```python
some_strings = np.array(['fruits', '🍎', '🍇', '🍓', 'mouse', '🐭', 'panda', '🐼'])
some_strings
```

What is the integer position of `'🍓'` in the array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Your answer should be a **positive** integer!

```python
strawberry_position = ...
strawberry_position
```

```python
grader.check("q5_4")
```

**Question 5.5.** Suppose you have an array with 249 elements. What is the integer position of the sixth-to-last element in this array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Again, your answer should be a **positive** integer!

```python
sixth_last_position = ...
sixth_last_position
```

```python
grader.check("q5_5")
```

**Question 5.6.** Suppose you have an array with 743 elements. At what integer position is the middle element of this array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Again, your answer should be a **positive** integer!

```python
mid_position = ...
mid_position
```

```python
grader.check("q5_6")
```

By the way, it's also possible to use negative integer positions to access elements in an array, which can be easier than using positive integer positions sometimes.  If a position is negative, you count from the end of the array rather than from the beginning. Position -1 corresponds to the last element, -2 corresponds to the second-to-last element, and so on. For instance, to find the third-to-last element of `some_strings`, we could use:

```python
some_strings[-3]
```

## 6. DSC 10 Enrollments 📈

In this question, we'll look at enrollment data from the fall and spring offerings of DSC 10 in the 2025-2026 academic year. The fourth column of the table below shows how many UCSD undergraduate students in the spring offering of DSC 10 come from each of the university's eight colleges. It looks like Revelle College is the most popular, with 27 students. The last column shows how many of those students are DSC majors. There are hardly any!

For comparison's sake, we also have the corresponding data from the Fall 2025 offering of DSC 10. You can see that most DSC majors take this course in the fall! 

Throughout this problem, we'll assume that all students in DSC 10 come from one of the eight colleges in the table. We've excluded graduate students, exchange students, etc.

| College   |   Fall 25 Students |   Fall 25 DSC Major Students |   Spring 26 Students |   Spring 26 DSC Major Students |
|:----------|-------------------:|-----------------------------:|---------------------:|-------------------------------:|
| Eighth    |                 43 |                           30 |                   8 |                              1 |
| ERC   |                 38 |                           20 |                   19 |                              0 |
| Muir     |                 54 |                           42 |                   13 |                              0 |
| Revelle |                 42 |                           27 |                   27 |                              0 |
| Sixth    |                 81 |                           59 |                   25 |                              2 |
| Seventh  |                 72 |                           46 |                   18 |                              0 |
| Marshall      |                 56 |                           38 |                   26 |                              0 |
| Warren   |                 67 |                           51 |                   14 |                              1 |

In this problem, we'll be working with the data from this table as *arrays*. Here are those arrays:

```python
students_25 = np.array([43, 38, 54, 42, 81, 72, 56, 67])
students_25
```

```python
majors_25 = np.array([30, 20, 42, 27, 59, 46, 38, 51])
majors_25
```

```python
students_26 = np.array([8, 19, 13, 27, 25, 18, 26, 14])
students_26
```

```python
majors_26 = np.array([1, 0, 0, 0, 2, 0, 0, 1])
majors_26
```

Remember, the `numpy` package (`np` for short) provides many handy functions for working with arrays. These are specifically designed to work with arrays and are faster than using Python's built-in functions. 

Some frequently used array functions are `np.min`, `np.max`, `np.sum`, `np.abs`, and `np.round`. There are many more, which you can browse by typing `np.` into a code cell and hitting the `tab` key. These are all in addition to array methods, which we access using dot notation. (For instance, both `arr.max()` and `np.max(arr)` return the largest element in the array `arr`.)

**Question 6.1.** Assign `enrolled_25` and `enrolled_26` to the number of students that were enrolled in DSC 10 in Fall 2025 and Spring 2026, respectively.

```python
enrolled_25 = ...
enrolled_26 = ...

# Don't change the lines below.
print('Students in Fall 2025:', enrolled_25)
print('Students in Spring 2026:', enrolled_26)
```

```python
grader.check("q6_1")
```

**Question 6.2.** How many non-DSC major students from the Fall 2025 offering of DSC 10 came from each of the eight colleges? Your answer should be an array called `non_majors_25`, with the colleges in the same order as they appear in the table above. For instance, the first element of `non_majors_25` should be the number of non-DSC majors in the Fall 2025 offering of DSC 10 who were in Eighth College.

```python
non_majors_25 = ...
non_majors_25
```

```python
grader.check("q6_2")
```

**Question 6.3.** What percentage of Fall 2025 DSC 10 students from each college were DSC majors? Your answer should be an array called `percent_majors_25`, with the colleges in the same order as they appear in the table above, and with percentages rounded to two decimal places.

```python
percent_majors_25 = ...
percent_majors_25
```

```python
grader.check("q6_3")
```

## 7. Analyzing Bob Ross Paintings 🎨

<center><img src="./images/bob_ross.jpg" width = 400/></center>

Bob Ross was an iconic American landscape painter, most well-known for hosting an instructional television show in the 80s and 90s called *The Joy of Painting*. In each episode, Bob Ross taught viewers to paint a different landscape picture while encouraging them to bring out "the artist hidden in the bottom of every single one of us." The file `data/bob_ross.csv` contains information on individual elements seen within Bob Ross' paintings across 403 episodes. For example, in [Season 1, Episode 1](https://www.youtube.com/watch?v=oh5p5f5_-7A), Bob shows how to create the painting below, titled "A Walk in the Woods".

&nbsp;
<center><img src="./images/bob_ross_s1_painting.png" width = 400/></center>

**Question 7.1.** Read this file into a DataFrame called `bob_ross`. 

```python
bob_ross = ...
bob_ross
```

```python
grader.check("q7_1")
```

This dataset contains 403 rows and 71 columns. Each column contains a different element that may or may not appear in a painting. The dataset uses the number 1 to mean that the element is included in the painting, and 0 if not. Notice, for example, that in Season 1, Episode 1, the dataset records a 1 for `'TREES'` and a 0 for `'WINDMILL'`, which corresponds to what we see in the painting shown above.

It's hard to tell which different painting elements are included in our dataset, because the DataFrame above displays with an ellipsis (...) truncating the dataset in between the columns `'BRIDGE'` and `'TOMB_FRAME'`. This is because the `babypandas` display settings default to showing a maximum of 20 columns, the first 10 and the last 10. 

If you want to see the names of all the columns, you can run the cell below. 

```python
bob_ross.columns
```

**Question 7.2.** Examine the columns `'EPISODE'`, `'SEASON'`, and `'EPISODE_NUMBER'` and consider which would be the best choice of index for this DataFrame. Change the `bob_ross` DataFrame so that it's indexed by this value instead of the default index.

```python
bob_ross = ...
bob_ross
```

```python
grader.check("q7_2")
```

***Note:*** Try running the cell where you set the index a second time. You should see an error message. Can you figure out what the error message means and why that's happening? If not, try asking Google or the AI Tutor, a great first step when you come across an error that you don't understand. If you're still confused, feel free to ask a classmate or the staff on Campuswire.

When you submit your work for autograding, the entire notebook will be run from start to finish. Each cell will run only once, so it's no problem if your code errors on the second run. In this case, it means you're doing something right!

**Question 7.3.** Throughout all 31 seasons, how many of Bob’s paintings contained an ocean? Assign your result to `ocean`. 

```python
ocean = ...
ocean
```

```python
grader.check("q7_3")
```

**Question 7.4.** What was the first episode in which a lighthouse was painted? Assign your result to `first_lighthouse` as a string in the form `'S__E__'`. For example, if your answer is Season 1 Episode 1 then your answer should be `'S01E01'`. 

```python
first_lighthouse = ...
first_lighthouse
```

```python
grader.check("q7_4")
```

**Question 7.5.** Which episode titles contain the word `'MOUNTAIN'`? Create an *array* called `mountain_episodes` containing the titles of all such episodes. Then assign to `num_mountain_episodes` the number of such episodes.

***Hint:*** To convert a Series into an array, call the function `np.array` on the Series.

```python
mountain_episodes = ...
num_mountain_episodes = ...
num_mountain_episodes
```

```python
grader.check("q7_5")
```

**Question 7.6.**  What proportion of all the paintings with a `'RIVER'` have both a `'WATERFALL'` and `'GRASS'`? Assign your result to the variable `river`. Remember, proportions are always between 0 and 1.

***Hint:*** You'll need to use multiple conditions in order to solve this problem. [This page](https://dsc-courses.github.io/bpd-reference/docs/documentation/querying/df[]) in the reference website will be very useful. 

```python
river = ...
river
```

```python
grader.check("q7_6")
```

**Question 7.7.** For the paintings in Season 1, calculate the average number of elements in each painting. Assign this average to the variable `season1_average`.

```python
season1_average = ...
season1_average
```

```python
grader.check("q7_7")
```

**Question 7.8.** Create an array called `ranked_seasons` containing the season numbers (1 through 31) ordered in descending order according to the average number of elements in paintings from that season. Each element of `ranked_seasons` should be an `int`.

***Hint:*** After doing some grouping, you'll want to grab the index of the DataFrame using `.index`, then convert it to an array by feeding it in as an input to `np.array`.

```python
ranked_seasons = ...
ranked_seasons
```

```python
grader.check("q7_8")
```

## Finish Line: Almost there, but make sure to follow the steps below to submit! 🏁

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

**Citations:** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please tell us which tools you used by setting `gen_ai` to a **list** of numbers. Include all tools used, from the options below:

1. AI Tutor, in Tutor Mode
2. AI Tutor, in ChatGPT Mode
3. Any other outside generative AI tool, such as Gemini, Claude, ChatGPT, etc.

If you didn't use any such tools, set `gen_ai` to an empty list, `[]`.

```python
gen_ai = ...
```

```python
grader.check("citations")
```

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

To submit your assignment:

1. Make sure to comment out the code in Question 2.1 that causes an error.
1. Select `Kernel -> Restart Kernel and Run All Cells` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook as a `.ipynb` file using `File -> Download`, then upload your notebook  to Gradescope under "Lab 0".
5. **Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.**
6. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

With homeworks, unlike with labs, the grade you see on Gradescope is **not your final score**. We will run correctness tests after the assignment's due date has passed.

```python
grader.check_all()
```

Tests for citations (points: None)
```
>>> isinstance(gen_ai, list)
True
```
failure_message: In Python, a list has square brackets and comma-separated entries, such as [1, 2, 3]


Tests for q1_1 (points: None)
```
>>> import numbers
>>> bool(isinstance(scarecrow_debut, numbers.Integral) & isinstance(tin_man_debut, numbers.Integral) & isinstance(lion_debut, numbers.Integral))
True
```
```
>>> bool(scarecrow_debut in range(1, 13) and tin_man_debut in range(1, 13) and (lion_debut in range(1, 13)))
True
```


Tests for q1_2 (points: None)
```
>>> import numbers
>>> bool(isinstance(lion_mentioned, numbers.Integral))
True
```
```
>>> bool(lion_mentioned in [1, 2, 3, 4, 5])
True
```


Tests for q1_3 (points: None)
```
>>> import numbers
>>> bool(isinstance(oz_part3, numbers.Integral))
True
```
```
>>> bool(oz_part3 in [1, 2, 3, 4])
True
```


Tests for q1_4 (points: None)
```
>>> import numbers
>>> bool(isinstance(oz_part4, numbers.Integral))
True
```
```
>>> bool(oz_part4 in [1, 2, 3, 4, 5])
True
```


Tests for q2_1 (points: None)
```
>>> import numbers
>>> bool(isinstance(basics_1, numbers.Integral))
True
```
```
>>> bool(basics_1 in [1, 2, 3, 4])
True
```


Tests for q2_2 (points: None)
```
>>> import numbers
>>> bool(isinstance(basics_2, numbers.Integral))
True
```
```
>>> bool(basics_2 in [1, 2, 3, 4])
True
```


Tests for q3_1 (points: None)
```
>>> bool(route_A > 0)
True
```
```
>>> bool(int(route_A * 10 ** 4 % 10 ** 4) == 8333)
True
```
failure_message: Try again!


Tests for q3_2 (points: None)
```
>>> bool(route_B > 0)
True
```
```
>>> bool(int(route_B * 10 ** 10 % 10 ** 6) == 4564)
True
```
failure_message: Try again!


Tests for q3_3 (points: None)
```
>>> bool(side_length > 0)
True
```
```
>>> bool(int(side_length * 10 ** 7 % 10 ** 3) == 278)
True
```
failure_message: Try again!


Tests for q4_1 (points: None)
```
>>> bool(0 < dissimilarity < 100)
True
```
```
>>> bool(int(dissimilarity ** 0.5 * 10 ** 6 % 10 ** 5) == 45751)
True
```
failure_message: Try again!


Tests for q4_2 (points: None)
```
>>> import numbers
>>> bool(isinstance(most_dissimilar, numbers.Integral))
True
```
```
>>> bool(most_dissimilar in [1, 2, 3])
True
```


Tests for q4_3 (points: None)
```
>>> import numbers
>>> bool(isinstance(disposable, numbers.Integral))
True
```
```
>>> bool(disposable in range(1, 4))
True
```


Tests for q5_1 (points: None)
```
>>> bool(isinstance(quirky_numbers, np.ndarray) and len(quirky_numbers) == 5)
True
```
```
>>> bool(int(quirky_numbers[0] * 10 ** 4 % 10 ** 3) == 853)
True
```
failure_message: The first element of your array is not correct. Try again!
```
>>> bool(int(quirky_numbers[1] * 10 ** 3 % 10 ** 3) == 575)
True
```
failure_message: The second element of your array is not correct. Try again!
```
>>> bool(int(quirky_numbers[2] * 10 % 10 ** 3) == 66)
True
```
failure_message: The third element of your array is not correct. Try again!
```
>>> bool(int(quirky_numbers[3] * 10 ** 3 % 10 ** 3) == 388)
True
```
failure_message: The fourth element of your array is not correct. Try again!
```
>>> bool(int(quirky_numbers[4] * 10 ** 3 % 10 ** 3) == 944)
True
```
failure_message: The fifth element of your array is not correct. Try again!


Tests for q5_2 (points: None)
```
>>> bool(isinstance(likes, np.ndarray))
True
```
```
>>> bool(likes[0] == 'I like surfing')
True
```


Tests for q5_3 (points: None)
```
>>> bool(isinstance(by_comma, str) and isinstance(by_space, str))
True
```
```
>>> bool(len(by_comma) == 36)
True
```
failure_message: You have the wrong length for by_comma. Make sure by_comma appears EXACTLY as described.


Tests for q5_4 (points: None)
```
>>> import numbers
>>> bool(isinstance(strawberry_position, numbers.Integral) and strawberry_position > 0)
True
```
failure_message: Use a positive integer position.


Tests for q5_5 (points: None)
```
>>> import numbers
>>> bool(isinstance(sixth_last_position, numbers.Integral) and sixth_last_position > 0)
True
```
failure_message: Use a positive integer position.


Tests for q5_6 (points: None)
```
>>> import numbers
>>> bool(isinstance(mid_position, numbers.Integral) and mid_position > 0)
True
```
failure_message: Use a positive integer position.


Tests for q6_1 (points: None)
```
>>> import numbers
>>> bool(isinstance(enrolled_25, numbers.Integral) and isinstance(enrolled_26, numbers.Integral) and (enrolled_25 > enrolled_26))
True
```
failure_message: Make sure the order is correct!
```
>>> bool(int(enrolled_25 / enrolled_26 * 10 ** 3 % 10 ** 2) == 20)
True
```
failure_message: Try again!


Tests for q6_2 (points: None)
```
>>> isinstance(non_majors_25, np.ndarray)
True
```
```
>>> bool(int(non_majors_25[1] * 11 ** 3 % 10 ** 2) == 58)
True
```
failure_message: Try again!
```
>>> bool(int(non_majors_25[7] * 11 ** 3 % 10 ** 2) == 96)
True
```
failure_message: Try again!


Tests for q6_3 (points: None)
```
>>> bool(np.all(percent_majors_25 < 1) == False)
True
```
failure_message: Make sure to use percentages and not proportions!
```
>>> bool(round(percent_majors_25[0], 2) == percent_majors_25[0])
True
```
failure_message: Read the instructions carefully!
```
>>> bool(int(percent_majors_25[7] * 11 ** 3 % 10 ** 2) == 15)
True
```
failure_message: Try again!


Tests for q7_1 (points: None)
```
>>> bool(isinstance(bob_ross, bpd.DataFrame))
True
```


Tests for q7_2 (points: None)
```
>>> bob_ross.shape == (403, 70)
True
```
```
>>> bool(bob_ross.iloc[15].loc['TITLE'] == '"EBONY SEA"')
True
```
failure_message: Try again!


Tests for q7_3 (points: None)
```
>>> import numbers
>>> bool(isinstance(ocean, numbers.Integral))
True
```
```
>>> bool(ocean < 100 and ocean > 0)
True
```
failure_message: Try again!


Tests for q7_4 (points: None)
```
>>> bool(isinstance(first_lighthouse, str))
True
```
```
>>> bool(len(first_lighthouse) == 6)
True
```
```
>>> bool(first_lighthouse[4] == '0')
True
```


Tests for q7_5 (points: None)
```
>>> import numbers
>>> bool(isinstance(num_mountain_episodes, numbers.Integral) and isinstance(mountain_episodes, np.ndarray) and (num_mountain_episodes > 0))
True
```
```
>>> bool('"AUTUMN MOUNTAINS"' in set(mountain_episodes))
True
```


Tests for q7_6 (points: None)
```
>>> bool(isinstance(river, float) and 0 <= river <= 1)
True
```
```
>>> bool(int(river * 10 ** 3 % 10 ** 2 * 10) == 269)
True
```
failure_message: Try again!


Tests for q7_7 (points: None)
```
>>> bool(isinstance(season1_average, float))
True
```
```
>>> bool(int(season1_average * 10 ** 3 % 10 ** 2) == 38)
True
```
failure_message: Try again!


Tests for q7_8 (points: None)
```
>>> bool(isinstance(ranked_seasons, np.ndarray))
True
```
```
>>> import numbers
>>> bool(all((isinstance(season, numbers.Integral) for season in ranked_seasons)))
True
```
```
>>> bool(len(ranked_seasons) == 31)
True
```
```
>>> bool(ranked_seasons[0] == 12 and ranked_seasons[1] == 10 or (ranked_seasons[0] == 10 and ranked_seasons[1] == 12))
True
```
failure_message: Try again!
